# Week 5 Problem Set

## Cohort

```python
# %load_ext nb_mypy
# %nb_mypy On
```

```python
from typing import TypeAlias
from typing import Optional, Any
from __future__ import annotations

Number: TypeAlias = int | float
```

**CS1.** Implement a `RacingGame` class that plays car racing using Python `random` module to simulate car's acceleration. The class has the following attribute(s):
- `car_list` which is a dictionary containing all the `RacingCar` objects where the keys are the racer's name.

The class has the following properties:
- `winners` which list the winners from the first to the last. If there is no winner, it should return `None`.

Upon instantiation, it should initalize the game with some **random seed**. This is to ensure that the behaviour can be predicted.

It has the following methods:
- `add_car(name, max_speed)` which creates a new `RacingCar` object and add it into the `car_list`. 
- `start(finish_distance)` which uses the `random` module to assign different initial speeds (0 to 50) to each of the racing car and set the same finish distance for all cars.
- `play(finish)` which contains the main loop of the game that calls the `RacingCar`'s method `race()` until all cars reach the finish line. It takes in an argument for the finish distance.



```python
class RacingCar:
    
    def __init__(self, name: str, max_speed: int) -> None:
        self.racer: str = name
        self.max_speed: int = max_speed
        self.finish: int = -1
        self.pos = 0
    
    @property
    def racer(self) -> str:
        return self._racer
    
    @racer.setter
    def racer(self, name: str) -> None:
        if isinstance(name, str) and name!= "":
            self._racer = name
            
    @property
    def speed(self) -> int:
        return self._max_speed
    
    @speed.setter
    def speed(self, val: int) -> None:
        if val >= 0 and val <= self.max_speed:
            self._max_speed = val
    
    @property
    def pos(self) -> int:
        return self._pos
    
    @pos.setter
    def pos(self, val: int) -> None:
        if val >=0:
            self._pos = val
            
    @property
    def is_finished(self) -> bool:
        if self.pos >= self.finish > 0:
            return True
        else: 
            return False
            
    def start(self, init_speed: int, finish_dist: int) -> None:
        if self.speed <= self.max_speed:
            self.speed = init_speed
        if finish_dist >= 0:
            self._pos = 0
            self.finish = finish_dist
    
    def race(self, acc: int) -> None:
        new_speed = self.speed + acc
        self.speed = new_speed
        self.pos = self._pos + new_speed
        
    def __str__(self) -> str:
        return f"Racing Car {self.racer} at position: {self.pos}, with speed: {self.speed}."
    
            
    
```

```python
import random

class RacingGame:
    
    def __init__(self, seed: Number) -> None:
        self.car_list: dict[str, RacingCar] = {}
        self._winners: list[str] = []
        random.seed(seed)
        
    @property
    def winners(self) -> Optional[list[str]]:
        if len(self._winners) != 0:
            return self._winners
        else:
            return None
         
    def add_car(self, name: str, speed: int) -> None:
        self.car_list[name] = RacingCar(name,speed) 
        
    def start(self, finish: int) -> None:
        for keys in self.car_list.keys():
            self.car_list[keys].speed = random.randint(0,50)
            self.car_list[keys].finish = finish
    
    def play(self, finish: int) -> None:
        self.start(finish)
        finished_car: int = 0
        while True:
            for racer, car in self.car_list.items():
                if not car.is_finished:
                    acc: int = random.randint(-10, 20)
                    car.race(acc)
                    # you can comment out the line below to check the output
                    # print(car)
                    if car.is_finished:
                        self._winners.append(racer)
                        finished_car +=1
            if finished_car == len(self.car_list):
                break
            
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
game: RacingGame = RacingGame(100)
assert game.car_list == {}
assert game.winners == None

game.add_car("Hamilton", 250)
assert len(game.car_list) == 1
assert game.car_list["Hamilton"].racer == "Hamilton"

game.add_car("Vettel", 200)
assert len(game.car_list) == 2
assert game.car_list["Vettel"].racer == "Vettel"

game.start(200)
assert [ car.pos for car in game.car_list.values()] == [0, 0]
assert [ car.speed for car in game.car_list.values()] == [9, 29]
assert [ car.finish for car in game.car_list.values()] == [200, 200]

game.play(200)
assert game.winners == ["Vettel", "Hamilton"]

game: RacingGame = RacingGame(200)
game.add_car("Hamilton", 250)
game.add_car("Vettel", 200)
game.play(200)
assert game.winners == ["Hamilton", "Vettel"]

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Write a class called `EvaluatePostfix` that evaluates postfix notation implemented using Stack data structure. Postfix notation is a way of writing expressions without parenthesis. For example, the expression `(1+2)*3` would be written as `1 2 + 3 *`. The class `EvaluatePostfix` has the following methods:
- `input(inp)`: which pushes the input one at a time. For example, to create a postfix notation `1 2 + 3 *`, we can call this method repetitively, e.g. `e.input('1'); e.input('2'); e.input('+'); e.input('3'); e.input('*')`. Notice that the input is of String data type. 
- `evaluate()`: which returns the output of the expression.

Postfix notation is evaluated using a Stack. The input streams from `input()` are stored in a Queue, which we will implement using Python's List. Note: If you have finished your homework on Queue, you can replace this part with your Queue. 

If the output of the Queue is a number, the item is pushed onto the stack. If it is an operator, we will apply the operator to the top two items in the stacks, pushing the result back onto the stack. 

```python
class Stack:
    def __init__(self) -> None:
        self.__items: list[Any] = []

    def push(self, item:Any):
        self.__items.append(item)

    def pop(self) -> Any:
        if len(self.__items) != 0:
            val = self.__items[-1]
            del self.__items[-1]
            return val
        else:
            print("nonetype")
            return None

    def peek(self) -> Any:
        if len(self.__items) != 0:
            return self.__items[-1]
        else:
            return None

    @property
    def is_empty(self) -> bool:
        if len(self.__items) > 0:
            return False
        else:
            return True

    @property
    def size(self):
        return len(self.__items)

```

```python
class EvaluatePostfix:

    operands: str = "0123456789"
    operators: str = "+-*/"

    def __init__(self) -> None:
        self.expression: list[str] = []
        self.stack: Stack = Stack()

    def input(self, item: str) -> None:
        if item in self.operands or item in self.operators:
            self.expression.append(item)

    def evaluate(self) -> Number:
        for i in range(len(self.expression)):
            if self.expression[i] in self.operands:
                self.stack.push(self.expression[i])
            else:
                rightly = self.stack.pop()
                statement = str(self.stack.pop())+ self.expression[i] +str(rightly)
                self.stack.push(eval(statement))
        return self.stack.peek()
                
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
pe: EvaluatePostfix = EvaluatePostfix()
pe.input("2")
pe.input("3")
pe.input("+")
assert pe.evaluate()== 5

pe.input("2")
pe.input("3")
pe.input("+")
pe.input("6")
pe.input("-")
assert pe.evaluate()== -1

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Inheritance:* Create a class called `MixedFraction` as a subclass of `Fraction`. A mixed fraction is a fraction that comprises of a whole number, a numerator and a denominator, e.g. `1 2/3` which is the same as `5/3`. The class has the following way of initializing its properties:
- `__init__(top, bot, whole)`: which takes in three Integers, the whole number, the numerator, and the denominator, e.g. `whole=1`, `top=2`, `bot=3`. The argument `whole` by default is `0`.  You can also specify `top` to be greater than `bot`. 

The class only has two properties:
- `num`: which is the numerator and can be greater than denominator.
- `den`: which is the denominator and must be a non-zero number.

The class should also have the following methods:
- `get_three_numbers()`: which is used to calculate the whole number, numerator and the denominator from a given numerator and denominator. The stored properties are `num` and `den` as in `Fraction` class. This function returns three Integers as a tuple, i.e. `(top, bot, whole)`.

The class should also override the `__str__()` method in this manner:
- `num/dem` if the numerator is smaller than the denominator. For example, `2/3`. 
- `whole top/bot` if the numerator is greater than the denominator. For example, `1 2/3`.

```python
def gcd(a: int, b: int) -> int:
    if b == 0:
        return a
    else:
        return gcd(b, a % b)


class Fraction:
    def __init__(self, num: int, den: int) -> None:
        self.num = num
        self.den = den
        pass
    
    @property
    def num(self) -> int:
        return self._num
    
    @num.setter
    def num(self, val: int) -> None:
        self._num = int(val)
    
    @property
    def den(self) -> int:
        return self._den
        pass
    
    @den.setter
    def den(self, val: int) -> None:
        if val == 0:
            self._den = 1
        else:
            self._den = val
    
    def __str__(self) -> str:
        return f"{self._num}/{self._den}"
    
    def simplify(self) -> Fraction:
        x = gcd(self._num, self._den)
        a , b = self._num/x , self._den/x
        return Fraction(int(a),int(b))
        
        
    
    def __add__(self, other) -> Fraction:
        x = self._den * other._den
        y = (other._num *self._den) + (self._num * other._den)
        return Fraction(y, x).simplify()

    
    def __sub__(self, other) -> Fraction:
        x = self._den * other._den
        y = (self._num * other._den) - (other._num *self._den) 
        return Fraction(y, x).simplify()
    
    def __mul__(self, other) -> Fraction:
        x = self._den * other._den
        y = (self._num * other._num)
        return Fraction(y, x).simplify()
    
    def __eq__(self, other) -> bool:
        if isinstance(other, Fraction):
            return self._num * other._den == self.den * other.num
        return False
    
    def __lt__(self, other) -> bool:
        return self._den * other._num > self._num * other._den
    
    def __le__(self, other) -> bool:
        return self._den * other._num >= self._num * other._den
    
    def __gt__(self, other) -> bool:
        return self._den * other._num < self._num * other._den
    
    def __ge__(self, other) -> bool:
        return self._den * other._num <= self._num * other._den
    

```

```python
class MixedFraction(Fraction):
    def __init__(self, top: int, bot: int, whole: int=0) -> None:
        num: Optional[int] = None
        num = whole * bot +top
        super().__init__(num, bot)

    def get_three_numbers(self) -> tuple[int, int, int]:
        whole: Optional[int] = None
        top: Optional[int] = None
        bot: Optional[int] = None
        
        whole = self.num //self.den
        top = self.num % self.den
        bot = self.den
        
        return (top, bot, whole)

    def __str__(self) -> str:
        (top,bot, whole) = self.get_three_numbers()
        if whole == 0:
            return super().__str__()
        else:
            return f"{whole:d} {top:d}/{bot:d}"
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
mf1: MixedFraction = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (2, 3, 1)
mf2: MixedFraction = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result: Fraction = mf1 + mf2
assert result.num == 10 and result.den == 3

assert mf1 == mf2

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
mf1: MixedFraction = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (2, 3, 1)
mf2: MixedFraction = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result: Fraction = mf1 + mf2
assert result.num == 10 and result.den == 3

result: Fraction = mf1 * mf2
assert result.num == 25 and result.den == 9

mf3: MixedFraction = MixedFraction(1, 2, 1)
result: Fraction = mf1 - mf3
assert result.num == 1 and result.den == 6

assert str(mf1) == "1 2/3"

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs33, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Create a class `Vertex` to represent a vertex in a graph. The class `Vertex` has the following attributes:
- `id_`: to identify each vertex. This is of String data type.
- `neighbours`: which is a Dictionary where the keys are the neighbouring `Vertex` object instances that are connected to the current Vertex and the values are the weights of the edge between the current Vertex and the neighbouring vertices. 

The class should also have the following methods:

- `__init__(self, id_)`: which is used to initialized the attribute `id_`. By default, `id_` is set to an empty String . The attribute `neighbours` is always set to an empty dictionary.
- `add_neighbour(self, nbr_vertex, weight)`: which adds a neighbouring Vertex to the current Vertex. The second argument provides the weight of the edge between the current Vertex and the newly added neighbouring Vertex. By default, `weight` is `0`.
- `get_neigbours(self)`: which returns all the Vertices connected to the current Vertex as a list. The elements of the output list are of `Vertex` object instances.
- `get_weight(self, neighbour)`: which returns the weight of the requested neighbour. It should return `None` if the requested neighbour is not found.
- `__eq__(self, other)`: which returns true if the id of the current vertex object is the same as the `other` vertex's id. 
- `__lt__(self, other)`: which returns true if the id of the current vertex object is less than the `other` vertex's id.
- `__hash__(self)`: which calls the `hash()` function on `id_` and returns it. This allows the object to be a dictionary key. This is provided for you.
- `__str__(self)`: This method should return the id of the current vertex and a list of `id_`s of the neighbouring vertices, like `Vertex 2 is connected to: 3, 4, 5` .

```python
class Vertex:
    def __init__(self, id_: str="") -> None:
        self.id_: str = id_
        self.neighbours: dict[Vertex, Number] = {}
    
    def add_neighbour(self, nbr_vertex: Vertex, weight: Number=0) -> None:
        self.neighbours[nbr_vertex] = weight
    
    def get_neighbours(self) -> list[Vertex]:
        return list(self.neighbours.keys())
    
    def get_weight(self, neighbour: Vertex) -> Optional[Number]:
        if neighbour in self.neighbours:
            return self.neighbours[neighbour]
        else: return None
    
    def __eq__(self, other) -> bool:
        return self.id_ == other.id_
    
    def __lt__(self, other) -> bool:
        return self.id_ < other.id_
    
    def __hash__(self) -> int:
        return hash(self.id_)
    
    def __str__(self) -> str:
        x = []
        for i in range(len(self.neighbours)):
            x.append(self.get_neighbours()[i].id_)  
        return f'Vertex {self.id_} is connected to: {", ".join(x)}'
    
v1: Vertex = Vertex("1")
v2: Vertex = Vertex("2")
v1.add_neighbour(v2)
v3: Vertex = Vertex("3")
v1.add_neighbour(v3, 3)

v4: Vertex = Vertex("4")
print("prin", v1.get_neighbours()[1].id_)
print(str(v1)) 
```

Output:
```
prin 3
Vertex 1 is connected to: 2, 3
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
v1: Vertex = Vertex("1")
assert v1.id_ == "1" and len(v1.neighbours) == 0
v2: Vertex = Vertex("2")
v1.add_neighbour(v2)
assert v1.get_neighbours()[0].id_ == "2" and v1.neighbours[v1.get_neighbours()[0]] == 0
v3: Vertex = Vertex("3")
v1.add_neighbour(v3, 3)
assert v1.get_weight(v3) == 3
v4: Vertex = Vertex("4")
assert v1.get_weight(v4) == None
assert v1 < v2
assert v1 != v2
assert str(v1) == "Vertex 1 is connected to: 2, 3"


```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** Create a class `Graph` to represent a Graph. The class has the following attribute:
- `vertices`: which is a *dictionary* of Vertices. The keys are the `id`s of the Vertices and the values are `Vertex` object instances.

The class has the following property:
- `num_vertices`: which is a *computed* property that returns the number of vertices in the graph.

The class also has the following methods:
- `__init__(self)`: which initializes the graph with an empty dictionary.
- `_create_vertex(self, id_)`: which creates a new `Vertex` object with a given `id_`. This method is never called directly and is only used by `add_vertex(id_)`.
- `add_vertex(self, id_)`: which creates a new `Vertex` object, adding it into the dictionary `vertices`. The argument `id_` is a String. This method should call `_create_vertex(id_)`.
- `get_vertex(self, id_)`: which returns the `Vertex` object instance of the requested `id_`. The method should return `None` if the requested `id_` cannot be found. The argument `id_` is a String.
- `add_edge(start_v, end_v)`: which creates an edge from one Vertex to another Vertex. The arguments are the `id_`s of the two vertices and are both Strings. If any of the starting or ending vertices is not in the graph, the function should call `add_vertex()` to create these vertices before creating an edge between the two vertices. 
- `get_neighbours(self, id_)`: which returns a list of `id_`s all the neighbouring vertices (of the specified Vertex `id_`). It should return an empty list `[]` if `id_` cannot be found. The argument `id_` is a String and the elements of the output list are of `str` data type. 
- `__contains__(self, id_)`: which returns either `True` or `False` depending on whether the graph contains the specified Vertex's `id_`. The argument `id_` is a String.

```python
class Graph:
    def __init__(self) -> None:
        self.vertices: dict[str, Vertex] = {} #dictionary containing the key and vertex object respectively
        
    def _create_vertex(self, id_: str) -> Vertex:
        #which creates a new Vertex object with a given id_. This method is never called directly and is only used by add_vertex(id_).
        return Vertex(id_)
    
    def add_vertex(self, id_: str) -> None:
        #which creates a new Vertex object, adding it into the dictionary vertices. The argument id_ is a String. This method should call _create_vertex(id_).
        self.vertices.update({id_: self._create_vertex(id_)})
    
    def get_vertex(self, id_: str) -> Optional[Vertex]:
        if id_ in self.vertices.keys():
            return self.vertices[id_]
        return None
    
    def add_edge(self, start_v: str, end_v: str, weight: Number=0) -> None:
        if start_v not in self.vertices.keys():
            self.add_vertex(start_v)
        if end_v not in self.vertices.keys():
            self.add_vertex(end_v)
        self.vertices[start_v].add_neighbour(self.get_vertex(end_v), weight)
            
    def add_edge_undirected(self, start_v: str, end_v: str, weight: Number=0) -> None:
        if start_v not in self.vertices:
            self._create_vertex(start_v)
        if end_v not in self.vertices:
            self._create_vertex(end_v)
            
        self.vertices[start_v].add_neighbour(self.vertices[end_v],weight)
        self.vertices[end_v].add_neighbour(self.vertices[start_v],weight)
        
    def get_neighbours(self, id_: str) -> list[str]:
        lister = self.vertices[id_].get_neighbours()
        attr = [o.id_ for o in lister]
        return sorted(attr)
    
    def __contains__(self, val: str) -> bool:
        return val in self.vertices.keys()
    
    def __iter__(self):
        for k,v in self.vertices.items():
            yield v 
        
    @property
    def num_vertices(self):
        return len(self.vertices)

```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
g: Graph = Graph()
assert g.vertices == {} and g.num_vertices == 0
g.add_vertex("A")
g.add_vertex("B")
g.add_vertex("C")
g.add_vertex("D")
g.add_vertex("E")
g.add_vertex("F")
assert g.num_vertices == 6
assert "A" in g
assert "B" in g
assert "C" in g
assert "D" in g
assert "E" in g
assert "F" in g
g.add_edge("A", "B")
g.add_edge("A", "C")
g.add_edge("B", "C")
g.add_edge("B", "D")
g.add_edge("C", "D")
g.add_edge("D", "C")
g.add_edge("E", "F")
g.add_edge("F", "C")
assert sorted(g.get_neighbours("A")) == ["B", "C"]
assert sorted(g.get_neighbours("B")) == ["C", "D"]
assert sorted(g.get_neighbours("C")) == ["D"]
assert [v.id_ for v in g] == ["A", "B", "C", "D", "E", "F"]

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Create a function to generate a `Graph` object instance for the following MRT lines. 

![](https://data-driven-world.github.io/2023/assets/images/MRT_Train-cd89106da00c927ac48ff3792b34ebc2.png)

You will have to use the following methods:
- `add_vertex()` to add a new vertex
- `add_edge()` to add a new edge between two vertices

Note: Since the image shows an undirected graph, you need to add an edge from vertex A to vertex B and another edge from vertex B to vertex A.

```python
stations = ["Bedok Reservoir", "Tampines West", 
            "Tampines", "Simei", "Tampines East", 
            "Pasir Ris", "Bedok", "Tanah Merah",
            "Upper Changi", 
            "Expo", "Changi Airport"]
```

```python
def create_sub_lines():
    g = Graph()
    for i in stations:
        g.add_vertex(i)
        
    

    g.add_edge_undirected("Bedok Reservoir","Tampines West")
    
    g.add_edge_undirected("Tampines West","Tampines")
    
    g.add_edge_undirected("Tampines","Pasir Ris")
    g.add_edge_undirected("Tampines","Simei")
    g.add_edge_undirected("Tampines","Tampines East")
    
    g.add_edge_undirected("Simei","Tanah Merah")
    
    g.add_edge_undirected("Tanah Merah","Bedok")
    
    g.add_edge_undirected("Tanah Merah","Expo")
    
    g.add_edge_undirected("Tampines East","Upper Changi")
    g.add_edge_undirected("Upper Changi","Expo")
    
    g.add_edge_undirected("Expo","Changi Airport")
    return g
```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
output = create_sub_lines()

assert output.num_vertices == 11
assert sorted(output.vertices) == sorted(stations)
assert output.get_neighbours("Bedok Reservoir") == ["Tampines West"]
assert output.get_neighbours("Tampines West") == ["Bedok Reservoir", "Tampines"]
assert sorted(output.get_neighbours("Expo")) == ["Changi Airport", "Tanah Merah", "Upper Changi"]

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```